---
id: 656469e62377ab34a5bcb8d1
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Declare two new variables named `digits` and `symbols` and assign them `string.digits` and `string.punctuation`, respectively.

# --hints--

Оголосіть змінну під назвою `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Призначте `string.digits` до змінної `digits`.

```js
({ test: () => assert(__pyodide.runPython(`
    import string
    nums = __locals.get("digits")
    nums == string.digits
  `))
})
```

Оголосіть змінну під назвою `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Призначте `string.punctuation` до змінної `symbols`.

```js
({ test: () => assert(__pyodide.runPython(`
    import string
    s = __locals.get("symbols")
    s == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
